## Write a haiku below (1 point)

<!-- grade_id: haiku-answer, points: 1, grade: true, solution: true -->
Haiku goes here

Another task.. worth two points, that is spread along multiple cells and manually graded.

<!-- grade_id: cell-ff14169315ad87d3, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n"""
    
    ### BEGIN SOLUTION
    if n < 1:
        raise ValueError("n must be greater than or equal to 1")
    return [i**2 for i in range(1, n+1)]
    ### END SOLUTION
```

<!-- grade_id: cell-dee87c78c1785f81, points: 3, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(squares(1), [1])
assert_equal(squares(2), [1, 4])

### BEGIN HIDDEN TESTS
assert_equal(squares(3), [1, 4, 9])
### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-4-bd297c1c2cbf> in <module>
----> 1 from nose.tools import assert_equal
      2 
      3 assert_equal(squares(1), [1])
      4 assert_equal(squares(5), [1])
      5 

ModuleNotFoundError: No module named 'nose'
```

<!-- grade_id: cell-36cd318c1c136ece, points: 4, grade: true, solution: false -->
```python
from nose.tools import assert_raises
assert_raises(ValueError, squares, 0)
assert_raises(ValueError, squares, -1)
```